```python
# Initialize Otter
import otter
grader = otter.Notebook("lab05.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Lab 6 – Visualization

## Data 6, Fall 2026
So far, we have discussed methods to interpret the data, but what if we want to present our data in a visual format? In this lab, you'll learn several important table methods for producing data visualizations. **Visualizations** are some of the most powerful tools in data science; they're helpful for showing data to people who don't necessarily have a background in data science, and allow data scientists like yourselves to help others understand the data in a more intuitive way.

We've talked about the `scatter` and `plot` methods, which allow us to visualize two numerical variables at once. We've also talked about the `barh` method, which allows us to display numerical values for categorical variables. These methods can open up more patterns in the data we otherwise would not have noticed.

As data scientists, it is not only our job to be able to implement various visualization methods, but also to know *when* to use each method. As we build our toolkit of visualization techniques going forward, it's important to understand the **advantages and disadvantages of each visualization type.**

```python
# Run this cell to load all required Python libraries
import numpy as np
from datascience import *

import seaborn as sns
import matplotlib.pyplot as plt
plt.style.use("fivethirtyeight")
%matplotlib inline
```

Something important to note before we begin is that the <code>salary</code> dataset that we'll be using today, which includes information on jobholders and their salaries, came from <a href=https://www.kaggle.com/datasets/mohithsairamreddy/salary-data/data>Kaggle</a> and was supposedly combined from multiple surveys, job postings, and other public sources. However, the Kaggle source does not provide any of the original sources that the data was taken from, so we have no idea how reliable or real this data is. It's okay to use data like this for the sake of practice, but when doing so, it is important to remember that the conclusions you can make become much less reliable and trustworthy. When looking to use data that can make an impact, be sure to thoroughly research where your data is coming from and how it was collected. Keep this fact in mind as you're going through the lab!

## Part 1: Scatter Plots

Let's begin by loading in the <code>unclean_salary_data.csv</code> file. This data required a lot of cleaning and manipulating behind the scenes in order to produce the visualizations in the lab, so be sure to keep that in mind: a lot of the time, you'll need to put work into preparing your data for analysis and visualization! Since some of the methods are out-of-scope for this course, we've done the cleaning for you beforehand.


```python
salary = Table.read_table("unclean_salary_data.csv")
salary.show(5)
```

### Contextualizing the Data
As data scientists, it's important to take a look at the data we're working with to understand the information we have available to us. Take some time to look at our `salary` table and try to understand what information we have.

<!-- BEGIN QUESTION -->

---
### Question 1a
What information does our table tell us? Additionally, what does each record (i.e. row) represent?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 1px solid #fdb515;" />

### **The [scatter](http://data8.org/datascience/_autosummary/datascience.tables.Table.scatter.html#datascience.tables.Table.scatter) method**

As we mentioned in lecture, visualizing variables can show us patterns in the data. The `scatter` method allows us to see the relationship between two numerical variables in our data by producing a **scatter plot**. The first provided column name goes along the x-axis and the second goes along the y-axis.

Let's take a look at the relationship between **years of experience** (`yr_exp`) and **salary** (`$@lary`) using our `salary` table.

### Producing Scatter Plots

Now, we can call `scatter` on the `salary` table. Run the following cell to do so.

```python
salary.scatter("yr_exp", "$@lary")
```

Just like that, you've produced your first scatter plot! It looks a little messy, however. While we know what the columns mean because we have access to the data, our audience does not know what `yr_exp` and `$@lary` are in our data. 

Let's try to clean up the plot to make it easier for our audience to understand.

---
### Question 1b 

Let's try to clean up the figure. We're going to keep the column names the same (for right now), but change the labels on our scatter plot. The label for your x-axis should be `Years of Experience`, and the y-axis label should be `Salary`. 

*Note*: Add `matplotlib.pyplot` functions after the line of code with your scatter plot to add labels.

```python
salary.scatter("yr_exp", "$@lary")
plt.title('CHANGE TITLE HERE')
...
```

```python
grader.check("q1_b")
```

---
### Question 1c

While that fixed the problem in our figure, it didn't fix the issue of badly worded columns in our table. Let's recreate the same plot, but now let's change the column names so our `salary` table is easier to use for future analyses.

*Hint*: Check your [Data 6 cheatsheet](https://data6.org/fa26/resources/cheatsheet/) for a function to rename column names. For simplicity, we're going to run the code twice to change the two column names we need to change. The new column labels will be `Years of Experience` and `Salary`.

```python
salary = ...
salary = ...

salary.scatter('Years of Experience', 'Salary')
```

```python
grader.check("q1_c")
```

Okay, now we have a more useful table and an interpretable scatter plot.

---
### Question 1d 

Finally, let's finalize our scatter plot so that our audience can understand what we are plotting. We're going to recreate the plot below, but let's also make the plot better by adding some additional information to our scatter plot, like a title.

![](/Users/jp/Desktop/github_projects/CDSS/assignments_new/notebooks/lab/lab05/figure1.png)

```python
salary_tick_values = ...
salary_tick_labels = ...

fig1_title = 'Relationship between years\nof experience and salaries'

salary.scatter('Years of Experience', 'Salary')
...
plt.yticks(salary_tick_values, labels = salary_tick_labels) 
```

```python
grader.check("q1d")
```

<!-- BEGIN QUESTION -->

---
### Question 1e

Often, scatter plots can suffer from what's known as **[overplotting](https://www.displayr.com/what-is-overplotting/)**: when many data points fall on top of each other, creating a blob of data. When this happens, it's often difficult to see the individual data points.

Why did we decide on a scatter plot a better rather than a line plot for comparing years of experience and salary?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---
### Question 1f

We can examine fewer data points by filtering rows to only include those with certain degrees.

We can filter our specific education levels; however, we don't know what levels of education we have. Let's get the counts of `Education Level` and assign the counts table to `education_counts`.

```python
education_counts = ...
education_counts
```

```python
grader.check("q1f")
```

---
### Question 1g

We are going to filter rows to only include those that have graduate degrees, so we'll filter for `Master's Degree` and `PhD` degrees. We'll create the `scatter_grad` table, where we will filter by the two graduate degrees. We will continue to plot `"Years of Experience"` on the x-axis and `"Salary"` on the y-axis.

*Hint*: Look at the `are.contained_in` documentation on the [Python Reference Sheet](https://data6.org/notes/reference#table.where-predicates) to help you out.

```python
scatter_grad = ...

scatter_grad.scatter('Years of Experience', 'Salary')
```

```python
grader.check("q1g")
```

---
### Question 1h

From the previous scatter plot, it seems like salary increases for individuals with graduate degrees similarly to other education levels.

Let's focus on those with the top half of years of experience. Specifically, let's focus on individuals with 15 and greater years of experience to see the scatter plot for these individuals. 

When filtering the data, we will be using `scatter_grad` from the previous question to further filter the data. 

```python
scatter_grad_most_experience = ...

scatter_grad_most_experience.scatter('Years of Experience', 'Salary')
```

```python
grader.check("q1h")
```

<!-- BEGIN QUESTION -->

---
### Question 1i

Since we filtered to include two different graduate degrees. We aren't sure what values belong to individuals with Master's degrees or those with PhDs. [We should look at the documentation for the scatter function](https://data8.org/datascience/_autosummary/datascience.tables.Table.scatter.html#datascience.tables.Table.scatter). This will give us an understanding of what additional arguments we can incorporate to see which points belong to individuals with Master's degrees and what points belong to those with PhDs. 

From earlier in this session, we also learned about the role of having 0 for an axis in our plots. Currently, our x-axis starts at 15. Since we filtered the data specifically to start at 15 years of experience, we shouldn't add a 0 to our x-axis. 

It is still our responsibility to be ethical data scientist so we will include information for our audience to know the scale of our data. One way of drawing attention to our x-axis' scale is to include a caption. [We can look at the documentation for how we can add a caption to our scatter plot](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.figtext.html#matplotlib.pyplot.figtext). 

We are going to create the same scatter plot as above, but this time we want to see if there is a difference between individuals with Master's degrees or PhDs in the relationship between `Years of Experience` and `Salary`. Additionally, we want to make sure we draw attention to our scale starting at 15 and not at zero.

*Hint*: Play around with the first two arguments for the caption to see what works best for your plot. There is no correct answer as long as the caption is able to be read and not overlapping with other components of the scatter plot.

```python
fig_caption = 'This plot\'s x-axis starts at 15 not 0.'

scatter_grad_most_experience.scatter('Years of Experience',
                                     'Salary', 
                                     ...,
                                     ...
                                     )
plt.figtext(..., ..., ...)
```

<!-- END QUESTION -->

<hr style="border: 1px solid #fdb515;" />

## Part 2

We're going to go back to using the `salary` table. This time we are going to focus on creating bar plots.

For our first task, let's get the counts of each individual based on their `Job Title` and their `Education Level`. Then we can gather the **top 10 jobs** so we could analyze the counts of these jobs.

```python
job_ed_counts = ...
job_ed_counts
```

```python
grader.check("q2")
```

<!-- BEGIN QUESTION -->

### Question 2b

The table we created is fine, but our audience may want to see this in a figure instead. Use the `bar` function on the `job_ed_counts` table to better see the counts for job titles.

Let's display bar plots for both the counts for `Job Title`.

```python
...


___.bar(..., ...)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2c

What is wrong with our plot? How can we alter the code to make the plot more interpretable for our audience?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2d

From your response to the previous question, how would you alter the code from Question 2b to better visualize the counts for `Job Title`.

```python
...
```

### Question 2e

From the top 10 careers, let's also count the number of education levels for these careers. We could technically count the values of `Education Level` from the previous table manually, but let's try to do this using Python.

How would we count the `Education Level` values from the `job_ed_counts` table?

```python
top10_ed_counts = ...
top10_ed_counts
```

### Question 2e continued

So we now have the top 10 job titles in our `job_ed_counts` table. Let's go back to the `salary` table and filter the job titles by these top 10 job titles. 

We are going to filter our data to these top 10 job titles without hard coding the values. So we are not going to use `make_array()` and write out the 10 job titles. Instead, we are going to filter by using `job_ed_counts.column('Job Title')`. 

Then we are going to get the average `Age`, `Years of Experience`, and `Salary` for the 10 job titles. **We don't want these averages for `Education Level` and `Gender`.** Finally, we'll sort by years of experience where our first row will be the lowest average number of years of experience.

```python
job_title_averages = ...
job_title_averages
```

```python
grader.check("q2e")
```

<!-- BEGIN QUESTION -->

### Question 2f

Let's plot each of the following relationships in the `job_title_averages` table.

- `Job Title` and `Age mean`
- `Job Title` and `Years of Experience mean`
- `Job Title` and `Salary mean`

```python
print(...)
print(...)
print(...)
```

<!-- END QUESTION -->

---

## Done! 🦖

---

## Pets of Data 6

Congrats on finishing lab!

<img src="paulina.JPG" width="50%" alt="photo of a dog jumping on another dog's head"/>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1_b (points: 1)
```
>>> all([i in salary.labels for i in salary.select('yr_exp', '$@lary').labels])
True
```


Tests for q1_c (points: 1)
```
>>> all([i in salary.labels for i in salary.select('Years of Experience', 'Salary').labels])
True
```


Tests for q1d (points: 1)
```
>>> expected_salary_tick_values = make_array(0, 50000, 100000, 150000, 200000, 250000)
>>> all(salary_tick_values) == all(expected_salary_tick_values)
True
```
```
>>> expected_salary_tick_labels = make_array('0', '50,000', '100,000', '150,000', '200,000', '250,000')
>>> all(salary_tick_labels) == all(expected_salary_tick_labels)
True
```


Tests for q1f (points: 1)
```
>>> expected_education_counts = salary.group('Education Level')
>>> expected_education_counts.where('Education Level', 'High School').column('count').item(0) == education_counts.where('Education Level', 'High School').column('count').item(0)
True
```


Tests for q1g (points: 1)
```
>>> expected_scatter_grad = salary.where('Education Level', are.contained_in(make_array("Master's Degree", 'PhD')))
>>> expected_scatter_grad.num_rows == scatter_grad.num_rows
True
```


Tests for q1h (points: 1)
```
>>> expected_scatter_grad_most_experience = scatter_grad.where('Years of Experience', are.above_or_equal_to(15))
>>> expected_scatter_grad_most_experience.num_rows == scatter_grad_most_experience.num_rows
True
```


Tests for q2 (points: 1)
```
>>> job_ed_counts.num_columns == 3 and job_ed_counts.num_rows == 10
True
```
```
>>> job_ed_counts.column('count').item(1) == 330
True
```
```
>>> job_ed_counts.column('Job Title').item(2) == 'Data Analyst'
True
```


Tests for q2e (points: 1)
```
>>> top10_ed_counts.where('Education Level', 'PhD').column('count').item(0) == 3
True
```
```
>>> expected_job_title_averages = salary.where('Job Title', are.contained_in(job_ed_counts.column('Job Title'))).drop(make_array('Education Level', 'Gender')).group('Job Title', np.mean).sort('Years of Experience mean')
>>> expected_job_title_averages.column('Job Title').item(0) == job_title_averages.column('Job Title').item(0) and expected_job_title_averages.column('Job Title').item(-1) == job_title_averages.column('Job Title').item(-1)
True
```
