```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Lab 2 - Python Names, Strings, & Arrays

## Data 6, Fall 2026

Welcome to the second lab section of Data 6!

In today's lab, we'll:

- Work through problems using print
- Review names and data types
- Practice working strings
- Use arrays for incidence and prevalence calculations
- Practice debugging throughout

This lab is due on **Friday, September 4th at 11:59PM**.

You must submit the assignment to [Gradescope](https://www.gradescope.com/).

```python
# Run this cell to load all required Python libraries 
import numpy as np
!pip install datascience
from datascience import *

%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use("fivethirtyeight")

import warnings
warnings.simplefilter('ignore')
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 1: Print

Try running this cell:

```python
print("Hello, World!")
```

And this one:

```python
print("\N{WAVING HAND SIGN}, \N{EARTH GLOBE ASIA-AUSTRALIA}!")
```

Every `print` expression prints a line. Run the next cell and notice the order of the output.

```python
print("First this line is printed,")
print("and then this one.")
```

<br/>

---

## Question 1.1: Print Practice


Given an array `arr`, write code that will display the first and last value of `arr` without [hardcoding](https://en.wikipedia.org/wiki/Hard_coding) the values of what is being printed out. Use negative indexing to get the last element.

*Hint*: Remember to reference your [course notes](https://data6.org/notes/03-arrays/#indexing).

*Hint*: Did you also know that there is [documentation](https://data8.org/datascience/reference-nb/datascience-reference.html#array.item()) for the library we are using in this class?

For example, given an `arr` with the values [1, 3, 5, 7], we should see the following:

```
1
7
```

```python
# Do not edit the following line of code
arr = make_array(1, 3, 5, 7)

# YOUR CODE HERE
...
...
```

---

## [Tutorial] The backslash character `\`


By now, you know that in Python strings can be surrounded by single quotes (`'`) or double quotes (`"`). But what if a string had both single quotes and double quotes? This happens a lot in text and dialogue:
                                                                            

```python
# uncomment the following cell to see the error.
# print("Nike's slogan is, "Just Do It."")
```

We can use a **backslash character** `\` to **escape** quote characters inside strings to keep it all part of one string. Below, `\"` is treated as a single character—the double-quote—instead of the end of a string.

```python
print("Nike's slogan is, \"Just Do It.\"")
```

```python
# and again, using single quotes
print('Nike\'s slogan is, "Just Do It."')
```

The backslash can also be used for special characters and typesetting. In the below cell, the `\n` is treated as a single special character: a **newline**.

```python
print("Hello\nworld!")
```

The backslash character has many uses! Besides strings, it can be used as a **line continuation character** to write a single line of Python on multiple lines, without changing execution:

```python
print("Hello" \
      "world")
```

---

## Question 1.2: Multiple arguments to print

Notice how we **typecast** `s` to be a string so we can concatenate the strings together!

```python
polygon = "square"
s = 4
print("The area of a " + polygon + \
      " with side length " + str(s) + \
      " is " + str(s ** 2) + ".")
```

You can also print the value of variables by including them as arguments, such as `name = "John"`, then `print("My name is", name)`, which would display "My name is John". The `print` function can be used for debugging purposes or to provide output to the user in interactive programs. It is a fundamental tool in Python for displaying information and verifying program behavior.

**Question**: Try the following exercise, where the goal is to **print** the sum of two numbers. Your output should read:

> "The sum of 5 and 3 is 8"

```python
# Exercise: Print the sum of two numbers
num1 = 5
num2 = 3
sum_result = num1 + num2
...
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 2: Names & Data Types

Let's review using names and figuring out the data types for each of the names we assign.

## Question 2: 

```python
days_in_week = 7
weeks_in_year = 52.14
months_in_year = 12
```

```python
# add to this code to be able to see all the types when this cell is run
type(days_in_week) 
type(weeks_in_year)
type(months_in_year)
```

Now that we have the types from our names, let's use the names to get the number of days in a year. We will call the new name `days_in_year`. 

What will the data type be for our `days_in_year` name if we use the `round()` function?


```python
days_in_year = ...
print(days_in_year)
type(days_in_year)
```

```python
rounded = ...
rounded
```

```python
grader.check("q2_1")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 3: Strings

## Question 3

Let's practice printing with strings now. Let's draft an email to someone to explain to them that there are two instructors for the class, and they both happen to be named Jo(h)n. The issue is that someone tried to give you some Python code to make it easier, but it has several errors. Fix the errors in the code below to draft the email that there are two separate instructors.

> Jon isn't a good name for our other instructor. It wouldn't make it easier to distinguish him from John so we'll just have to continue calling him JP.

*Hint*: Don't worry about putting in backslashes to separate the print statement to separate lines. You can keep your answer as one long print statement.

```python
instructor 1 = 'John'
instructor2 = 'Jon'

print(Instructor2, 'isn't a good name for our other instructor. It wouldn't make it easier to distinguish him from', instructor1, 'so we'll just have to continue calling him JP.')
```

```python
grader.check("q3_1")
```

## Question 3.2: f Strings

Now, if someone just met the two instructors from talking to them, but they both introduced themselves as Jon/John. One way to tell the difference is to count the number of letters in their names. How would we do that in Python?

Once you get the number of letters in their names, try recreating the text below with a print statment and f strings. 

> Jon actually likes to go by JP, but he goes by Jon with family. We also know that John is our lecturing instructor because his name is 4 characters long.

```python
john_length = ...
jp_length = ...

print(john_length)
print(jp_length)
```

```python
...
```

```python
grader.check("q3_2")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 4: Arrays

## Question 4

We have already have some practice calculating incidence in historic data; however, let's get some data that may be more recent. 

We all heard and read about the Cyclospora outbreak where we had a recall on lettuce across several states in the United States. So let's look at a week's worth of data on the number of sick people. Let's focus on the difference between incidence and prevalence, where incidence is the number of new cases, while prevalence focuses on the total number of sick people during a period of time. 

[Find where we gathered the data from](https://www.cdc.gov/cyclosporiasis/outbreaks/07-26/timeline.html)


| Illness Onset Date | Number of Sick People |
| --- | --- |
| 7/7/2026 | 663 |
| 7/8/2026 | 875 |
| 7/9/2026 | 934 |
| 7/10/2026 | 1001 |
| 7/11/2026 | 723 |
| 7/12/2026 | 520 |
| 7/13/2026 | 370 |

$$\text{Incidence} = \frac{\text{new cases}}{\text{population at risk}}\;\times\;100,000$$

$$\text{Prevalence} = \frac{\text{total cases}}{\text{population at risk during this week}}\;\times\;100,000$$

```python
# don't edit this code
us_population = 342818000 
```

```python
new_sick_people = ...
```

```python
incidence = ...
print(incidence)
```

```python
total_sick_people = ...
prevalence = ...
print(prevalence)
```

```python
grader.check("q4_1")
```

## Question 4.2: Percentage Change

Let's find the **percentage change in the number of sick people between 7/10 and 7/11**. Since we are interested in the percentage change between the number of sick people, we will be using the raw number of sick people for both days. Rather than hardcoding the values, we are going to extract the values from our `new_sick_people` array to make the calculation.

*Hint*: We'll use the `.item()` method to get the number of sick people for both dates.

$$\text{percentage change} = \frac{B - A}{A} \times 100$$


```python
num_sick_people710 = ...
num_sick_people711 = ...

num_sick_people_diff = ...
round(num_sick_people_diff, 2)
```

```python
grader.check("q4_2")
```

## Question 5: Fixing JP's errors

Now that we have had some practice using names. Let's look at JP's code below and fix it. As an avid R user, he is prone to using the wrong arithmetic operator when he is trying to use exponentiation.

```python
x = '10'
x_squared = x^2
x_total = ...
```

```python
grader.check("q5_1")
```

## Question 5.2: Fixing More Errors

```python
item_prices = make_array(15.99, 23.49, 9.99, 5.49)
tax_rate = 8.25
discount = 10

total_price = sum(item_prices)
total_with_tax = total_price * tax_rate
total_with_discount = total_with_tax - (total_with_tax * (discount / 100))

print("Total price of items:", total_price)
print("Total after tax:", total_with_tax)
print("Total after discount:", total_with_discount)
```

## Pets of Data 6

Tigger is ready for the extended weekend after submitting that lab. Happy Friday!

<img src="tigger.jpg" width="50%" alt="photo of a cat sitting like a human on a bed"/>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q2_1 (points: 0)
```
>>> days_in_year == weeks_in_year * days_in_week
True
```
```
>>> rounded == type(round(days_in_year))
True
```


Tests for q3_1 (points: 0)
```
>>> instructor1 == 'John'
True
```
```
>>> instructor2 == 'Jon'
True
```


Tests for q3_2 (points: 0)
```
>>> john_length == len(instructor1)
True
```
```
>>> jp_length == len(instructor2)
True
```


Tests for q4_1 (points: 0)
```
>>> type(new_sick_people) is np.ndarray
True
```
```
>>> expected_sick_people = make_array(663, 875, 934, 1001, 723, 520, 370)
>>> (new_sick_people == expected_sick_people).all()
True
```
```
>>> (incidence == expected_sick_people / us_population * 100000).all()
True
```
```
>>> expected_total_sick_people = sum(expected_sick_people)
>>> (prevalence == expected_total_sick_people / us_population * 100000).all()
True
```


Tests for q4_2 (points: 0)
```
>>> num_sick_people710 == new_sick_people.item(3)
True
```
```
>>> num_sick_people711 == new_sick_people.item(4)
True
```
```
>>> num_sick_people_diff == 100 * (num_sick_people711 - num_sick_people710) / num_sick_people710
True
```


Tests for q5_1 (points: 0)
```
>>> expected_x = 10
>>> expected_x_squared = x ** 2
>>> x_total == expected_x + expected_x_squared
```
