```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Lab 4 – Tables and Data

## Data 6, Fall 2026

In this lab, we will be talking all about **tables**. We use tables to store all sorts of data form sports statistics to population information. If there's data you have ever been curious about, it is very likely that the internet has a table somewhere with that data!

Tables are integral to the foundation of Data Science, and in this lab we will go over how to query a table. Querying a table is basically asking information about the table. Some examples of common queries (in English, not code):

* How many data points are there?
* Which data points have a specific characteristic?
* What is the attribute of a specific data point?
* And many more!

This lab is due on **Sunday, September 20th at 11:59PM**.

You must submit the assignment to [Gradescope](https://www.gradescope.com/).

```python
# Run this cell to load all required Python libraries 
import numpy as np
from datascience import *

%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use("fivethirtyeight")

import warnings
warnings.simplefilter('ignore')
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 1: `None`

By now, you have seen that in many cases, evaluating Python expressions (function calls, etc.) returns values. But what if Python doesn't return _anything_? Python has a special value and type to represent this case: `None`.

In particular, the `print` returns `None`. How can that be, when we clearly see displayed output to the notebook???

Consider the following cell:

```python
# just run this cell
roller = "\N{ROLLER COASTER}"
print(roller)
roller + "\N{ROLLER SKATE}"
```

Above, notice that `print` was not the last line, yet still output to the screen. The last line was then evaluated and considered the output of the cell. You can verify this by the square brackets to the left of the output, compared to no square brackets next to the `print` output.

Now, considering the following more tricky code cell. Uncomment to run:

```python
# proclamation = print("We love " + roller)
# print("Your course staff says:", proclamation)
# proclamation + " and \N{ROLLER SKATE}"
```

What happened? **Discuss** your answers to the following questions (no text entry):
1. Which expression displays `We love 🎢` to the notebook?
1. What value is `proclamation` assigned to?
1. What happens in the second line?
1. Which line is erroring? _Hint_: string concatenation is not compatible when one of the values is `None` (which is not a string).

## Question 1

Let's fix the error you saw above. Fill in the below code so that when run, the cell prints:

```
We love 🎢
Your course staff says: We love 🎢
```

and has the below string as cell output:

```
'We love 🎢 and 🛼'
```

```python
proclamation = ...
...
...

# do not edit the line below
proclamation + " and \N{ROLLER SKATE}"
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 2: Unemployment Rates Introduction

The United States Bureau of Labor Statistics (BLS) publishes data about jobs in the US: [employment](https://www.bls.gov/charts/employment-situation/civilian-unemployment-rate.htm), [job openings, turnovers](https://www.bls.gov/jlt/jltover.htm), and so on.

Generally, the Bureau of Labor Statistics considers someone **unemployed** if that person does not have a job, has actively looked for work in the prior 4 weeks, and are currently available for work. The **unemployment rate** is the number of unemployed people as a percentage of the labor force.
    
Run the next cell to load the `unemployment_rates` table, taken from the [BLS](https://www.bls.gov/charts/employment-situation/civilian-unemployment-rate.htm). There will be no output—no output is expected as the cell contains an assignment statement. An assignment statement does not produce any output (it does not yield any value).

```python
# just run this cell
unemployment_rates = Table.read_table("unemployment.csv")
```

Let's examine our table to see what data it contains.

---

## Question 2.1: `show`

Use the method `show` to display the first 5 rows of `unemployment_rates`.

**Note**: The terms "method" and "function" are technically not the same thing, but for the purposes of this course, we will use them interchangeably.

**Hint**: `tbl.show(3)` will show the first 3 rows of the table named `tbl`. Additionally, make sure not to call `.show()` without an argument. This may crash your kernel! By argument, we mean the value within `.show()`. Looking at the [documentation](https://data8.org/datascience/_autosummary/datascience.tables.Table.show.html#datascience.tables.Table.show) we can see the argument for is `max_rows`, so we can print two rows of the table using `.show(max_rows = 2)`, we just don't always need to declare the argument.

```python
...
```

---

## Table attributes: `num_rows` and `num_columns`

We can ask for all sorts of information about the table itself:

```python
unemployment_rates.num_rows
```

```python
unemployment_rates.num_columns
```

---

## Question 2.2: `select`

Most of the columns are demographic-specific. If we're not interested in that information, it just makes the table difficult to read. This comes up more than you might think, because people who collect and publish data may not know ahead of time what people will want to do with it.

In such situations, we can use the table method select to choose only the columns that we want in a particular table. It takes any number of arguments. Each should be the name of a column in the table. It returns a new table with only those columns in it. The columns are in the order in which they were listed as arguments.

For example, the value of `unemployment_rates.select("Men, 20 years and over", "Women, 20 years and over")` is a table with only the unemployment rates of men and women 20 years and over. (However, this "drops" the critical time-related Month and Year columns).

**Question**: Use `select` to create a table with only the month, year, and total nationwide unemployment rate for that month. Call that new table `unemployment_totals`.

```python
unemployment_totals = ...
unemployment_totals
```

```python
grader.check("q2_2")
```

---

## Question 2.3: `drop`

`drop` serves the same purpose as `select`, but it takes away the columns that you provide rather than the ones that you don't provide. Like `select`, `drop` returns a new table.

**Question**: Suppose you didn't want the rates of the total unemployment, nor the men- and women-specific rates, nor the age-specific rates—just the rates by race and ethnicity. **Use `drop` to create a table that <u>doesn't include these four columns</u>.** Call that table `unemployment_races_ethnicities`.

```python
unemployment_races_ethnicities = ...
```

```python
grader.check("q2_3")
```

---

## Question 2.4: `sort`

Some details about `sort`:

* The first argument to `sort` is the name of a column to sort by.
* If the column has text in it, `sort` will sort alphabetically; if the column has numbers, it will sort numerically—both in ascending order by default.
* "Ascending order" means the values in the column will *increase* as you move *down the table*.
* `unemployment_totals.sort(label)` is a copy of `unemployment_totals` but sorted; the `unemployment_totals` table doesn't get modified.
* Rows always "stick together" when a table is sorted. It wouldn't make sense to sort just one column and leave the other columns alone. For example, if we sorted just the "Total" array, the total unemployment rates would all end up with the wrong months and years.

**Question**: Create a version of `unemployment_rates` that's sorted by total unemployment rate ("Total"), with the largest unemployment rate **first**. Call this new table `unemployment_rate_highest`.

_Hint_: To sort in descending order, you can provide the optional argument `descending=True`. See the [Python reference](https://data6.org/notes/reference).

```python
unemployment_rate_highest = ...
unemployment_rate_highest
```

```python
grader.check("q2_4")
```

**Discussion (no text entry)**: What month and year had the highest total unemployment rate? What happened globally in that timeframe?

---

## Question 2.5: `where`

We can use `where` to filter our table and look at specific rows. 

`where` takes 2 arguments:
1. The **name** of the column being filtered.
1. A **predicate** that describes the criterion that the column's values need to meet.

For now, we consider an "exact match" predicate. So `where` will return a copy of the original table, but with only rows where the specified column (first argument) is exactly the specified value (second argument).

**Question**: Use `unemployment_rates` to create a table called `unemployment_2025` containing unemployment rates for just `2025`. 

*Note* that our data doesn't include all 12 months for the year 2025.

```python
unemployment_2025 = ...
unemployment_2025
```

```python
grader.check("q2_5")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 3: Unemployment over Time

This question is more exploratory. A **time series** of data is sequential measurements of the same variable*, taken at regular intervals over time. (*We'll discuss what "variable" means in a social science context very soon.)

The data in `unemployment_rates` is a time series. We would like to visualize a two-dimensional line plot of the total national unemployment rate over time. However, if we use `unemployment_rates` to try to plot `"Total"` (on the y-axis) by `"Month"` (on the x-axis), we get a bizarre plot:

```python
# just run this cell
# we will learn much more about visualization syntax
# in the next unit.
unemployment_rates.plot("Month", "Total")
```

What's happening here? Simply put, the above plot did exactly what we wanted: It took the "Month" and "Total" values of each record in the table and plotted it, then connected the dots together. But two August unemployment rates can correspond to entirely different years (and therefore different times)! However, they would be plotted at the same horizontal value (here, August is 8). 

This issue happens because months are cyclic—meaning that once we hit December (12) of one year, the next month is January (1) of the next year. 

In this question, we seek to correct this visualization by using arrays to create a new strictly _increasing_ measure of time.

## Question 3.1: `column`

Before, we used `select` to create tables with a specific set of columns. Sometimes we may want to get a specific column as an array, so that we can perform additional processing.

The `column` method takes a column label or index and returns a copy of the values of that column as a NumPy array. For example, the below cell copies the "Total" column in `unemployment_rates` and stores it in the array `totals`:

```python
# just run this cell
totals = unemployment_rates.column("Total")
type(totals)
```

**Question**: Create two arrays `months` and `years` that are the columns of "Month" and "Year", respectively, in `unemployment_rates`.

```python
months = ...
years = ...
Table().with_columns("Month", months, "Year", years)
```

(Above, the last line ["pretty prints"](https://en.wikipedia.org/wiki/Pretty-printing) your results in table format, making it easier to visually check if you've done things correctly. See if you can understand what it's doing!)

```python
grader.check("q3_1")
```

---

## Question 3.2

Let's use the `months` and `years` arrays to construct a new array that accurately captures the time that the record was taken.

We'll do this by constructing a measure of time as fractions that combines months and years. We'll say that the first month (January) of 2006 is `2006.0`, the next month (February) is 2006 and 1/12 (or `2006.083333...`), the next month (March) is 2006 and 1/6 (or `2006.166666...`), and so on. This will give us 12 equally-spaced floating point values for each year, one corresponding to each month, chronologically sorted.

**Question**: Use `months` and `years` to create an array of times that follows this pattern and assign it to `times`. Your answer should involve array arithmetic.

_Hints_: This question is challenging! To help you debug:

1. The first month of every year, January, maps to a whole number (e.g., `2006.0`).
1. Like before, the last line "pretty prints" so you can check your work.

```python
times = ...
Table().with_columns("Month", months, "Year", years, "Time", times)
```

```python
grader.check("q3_2")
```

---

## Question 3.3: `with_columns`

We can create new tables by first making an empty `Table`, then calling `with_columns` on that table.

The `with_column` method requires two arguments:

1. The name of the column as a string
1. An array of values to put in the column

For example, `my_tbl.with_columns("My New Column", my_array)` would add to the table `my_tbl` a new column labeled "My New Column" with the values in `my_array`.

As a second example, `Table().with_columns("My New Column", my_array)` creates a one-column table.

**Question**: Create a new table with two columns labeled "Time" (using `times`, the fractional time you computed in the previous question), and "Total Unemployment" (using `totals`, the total unemployment rate per month). Call this new table `unemployment_over_time`.

```python
unemployment_over_time = ...
unemployment_over_time
```

```python
grader.check("q3_3")
```

---

## Question 3.4

Finally, let's use this new table to visualize our time series:

```python
# just run this cell
# we will learn much more about visualization syntax
# in the next unit.
unemployment_over_time.plot("Time", "Total Unemployment")
```

That looks much better! Notice that the visual spike in unemployment rate matches your conclusion from the tabular sorting you did earlier in lab.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Let's Get More Practice Working With Tables

<div class="alert alert-warning">
The data that we'll be using in this lab contains data on <a href=https://ir.ucr.edu/stats/admission/undergraduate>University of California, Riverside's undergraduate admissions.</a> We'll be looking at this data today as we are interested in exploring the context and ethics behind education-related data (UC Riverside specifically was chosen as it was the UC that provided the most publicly accessible method of downloading the data). In this case, we are looking at the demographics of undergraduate students. This leads us to many important questions to consider: why does UC Riverside make this data publicly available? Why does it choose this format and these categories to represent students? Who or what is affected by the conclusions we make from analyzing this data? Think about these questions as you work through the lab.
</div>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Part 4: Basic Table Functions

<div class="alert alert-warning">
In the first part of this lab, we'll be learning about some of the various basic table functions that we can utilize in the <code>datascience</code> library in order to answer questions about student demographics at UC Riverside. In this part, we'll be focusing on the demographic categories of <code>Gender</code> and <code>IPEDS Race/Ethnicity</code> to explore what the breakdown of students at UC Riverside looks like, as well as getting a better understanding of how and why UC Riverside chose to represent their data publicly in this format.
</div>

### Loading a Table

Recall in the previous lab, we introduced the `Table.read_table` method, which takes a *file path* and constructs a `Table` with the information from that file. Let's see how this works using the file `"riverside.csv"`, which contains information about UC Riverside's admissions.

*Note*: If you want to check where the `riverside.csv` file is, you can look in your DataHub directory by clicking `File` > `Open...` in the top left.

```python
ucr = Table.read_table("riverside.csv")
ucr.show(5)
```

<!-- BEGIN QUESTION -->

### Question 4.1 (Discussion)
As we discussed in the previous lab, what is the unit of analysis of our dataset? Another common way of asking this when considering our data is: what is the granularity of our data?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 4.2
Now that we have considered what each row in the data represents, let's take a look at how many rows we have overall. This gives us a better idea of approximately how many years and stratifications of students per year we have in our data. Assign the variable `ucr_rows` to the number of rows in `ucr`. 

You should not write an integer, but instead use one of the table attributes we have talked about so far to **calculate** the number of rows.

*Hint:* Stuck? Remember, you can reference all of the Table tools in `datascience` by looking at the Data 6 Python Reference sheet [here](http://data6.org/su25/reference). 

```python
ucr_rows = ...
ucr_rows
```

```python
grader.check("q4_2")
```

<div class="alert alert-warning">
To get our analysis started, let's take a look at the <code>drop</code> and <code>relabeled</code> methods to see how we can clean up the appearance of the columns and tables to make viewing and working with the data easier.
</div>

---
### Excluding columns: `drop`

We now have information about UC Riverside's admissions in a `datascience` Table. It is common in data analysis for our tables to have columns that may not be relevant to what our exploration is focused on. One way of cleaning up our table to remove these unnecessary columns is to use the `drop` method!

Along with having columns that may not be relevant to our exploration, it is very common for our datasets to have many rows or columns with `NaN` values (`nan` stands for "Not A Number", and it is typically used to denote a missing value). **However**, it is not a good idea to blindly drop columns with several NaN values from a table. Think back to what you saw with the missing values in Lab 2. What information would have been lost if we just dropped all missing values?

Our table below fortunately doesn't have any missing values, so we'll be using the `drop` method to get rid of columns that aren't relevant in our analysis. This will make it easier to focus on the columns that we do care about. Below, let's use the `drop` method to remove a column from our table, such as the `"New/Continuing Status"` column:

```python
ucr = ucr.drop("New/Continuing Status")
ucr.show(5)
```

*Note:* By dropping the `"New/Continuing Status"` column above, we no longer have information on whether the students in the stratification are new or continuing, so there may be some rows that appear to be duplicates. This is one of the downsides of dropping columns we aren't focusing on -- we may lose some data and context in the process! However, for our exploration's purposes, this information isn't that important, so we'll continue without it.

### Question 4.3 (Multiple Choice)
If the `"Full-Time Status"` column is dropped using the `drop` method, what would `ucr_dropped` look like? You can answer the question by assigning `q4_3` to `make_array(...)` where `...` is the choice of your answer (eg. `make_array(3)`).

1. It contains only rows with NaN values.
2. It contains only the columns before the `"Full-Time Status"` column.
3. It contains all columns except the `"Full-Time Status"` column.
4. It contains the same columns as the original table, but all NaN values in the `"Full-Time Status"` column are dropped.

```python
q4_3 = ...
```

```python
grader.check("q4_3")
```

---
### Changing column labels: `relabeled`

Sometimes the original column names of the data are long or not descriptive enough. We can rename column labels using the `relabeled` method. With this function, you are able to:
1. Relabel a *single column*
2. Relabel *several columns* at once

To change the names of multiple columns, we pass in an array of the old names and an array of the new names as the 2 inputs to `relabeled`.

*Note*: You may see another method called `relabel` in the `datascience` documentation. Please avoid using this, as it can change your data when you may not want to.

### Question 4.4
If we're accessing the `"IPEDS Race/Ethnicity"` column a lot, it could be cumbersome to type the long column name multiple times. Let's change the column label instead to `"Race/Ethnicity"`.

*Hint*: We've provided skeleton code for you to use.

```python
ucr_relabeled = ucr.relabeled(..., ...)
ucr_relabeled.show(5)
```

```python
grader.check("q4_4")
```

<div class="alert alert-warning">
In the next two sections of this part, we'll be looking at the <code>column</code> and <code>select</code> methods to discover ways of focusing on certain subsections of columns in our data.
</div>

---
### Querying the table: `column`

Looking at the first five rows of `ucr_relabeled` above, we can see some of the possible values for each of the columns. For example, in the `"Race/Ethnicity"` column, we see that there are values such as "Asian", "Black/African American", etc. In order to get a better idea of the individual values that exist within a column of our dataset, we can query the table using the `column` method. Let's use this method on `"Race/Ethnicity"`.

```python
ethnicity_list = ucr_relabeled.column("Race/Ethnicity")
ethnicity_list
```

One downside of doing it this way is that the list would be long and repetitive, since it will include the race/ethnicity for every row, resulting in repeated values. To work around this, we can use the `np.unique` method to tell us all the race/ethnicity values only once as they appear:

```python
np.unique(ethnicity_list)
```

<div class="alert alert-warning">
This gives us a much clearer picture of what the possible representations of race/ethnicity are in this dataset. Looking at the values above, here's something to consider: do you find any issues with how the individuals can be categorized? Are there any labels that could be misleading or misrepresentative? Are there any groups of people who might not be included by this breakdown? <br><br>

We'll consider this idea of representation of the data as we continue through the lab.
</div>

---
### Picking columns: `select`

In many cases during data analysis, there will be several columns that we aren't interested in, and only a few that we are. Instead of dropping several columns as we learned about earlier, we can use the `select` method to grab only the columns we want. 

### Question 4.5
For this example, we only want to keep the `"Year"`, `"Level"`, `"Full-Time Status"`, `"Race/Ethnicity"`, and `"Fall Headcount"`  columns. Fill in the following code so that the `ucr_relabeled` table has only the relevant columns.

```python
ucr_selected = ...
ucr_selected
```

```python
grader.check("q4_5")
```

### Question 4.6 (Multiple Choice)
What is the purpose of using the `select` method in this context? You can answer the question by assigning `q1_6` to `make_array(...)` where `...` is the choice of your answer (eg. `make_array(3)`).

1. To drop specific columns from the table selectively.
2. To filter rows based on a corresponding condition to keep relevant information.
3. To rearrange and rename columns in the table based on a criteria.
4. To choose specific columns to keep in the table and create a new table with only the selected columns.

```python
q4_6 = ...
```

```python
grader.check("q4_6")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Part 5: More Table Functions

<div class="alert alert-warning">
In Part 4 of this lab, we learned about several Table methods that help clean up and reorganize the appearance of a table to make our analysis easier. In this part, we'll learn about a couple more Table methods that help us filter and hone in on specific aspects of our data that we're interested in. In our case, these methods will help us better explore the breakdown of the gender and race/ethnicity demographics of students at UC Riverside.
</div>

### Sorting a column: `sort`

We will now introduce a new table method: `sort`. The `sort` table method allows us to see a table's column values sorted by its values in either **decreasing** (`descending=True`) or **increasing** (`descending=False`) order.

Let's say we want to ask the question: among all years, and particularly among the different groups of students based on their demographics, which fall headcount was the highest? Sorting the table will not only allow us to see a potential trend of headcounts over the years, but it will also show us if there are any trends among which groups of students have the highest headcounts. Let's use the `sort` method to find out:

```python
ucr_relabeled.sort("Fall Headcount", descending=True)
```

<div class="alert alert-warning">
The results that we see above are very interesting: firstly, we notice that the highest headcounts have all been in the most recent 10 years of the dataset (2013-2022), which shows us that student admissions at UC Riverside are higher than they were in the past; and secondly, we notice that the gender and race/ethnicity category is the same for the top 10 highest headcounts, female-identifying Chicanx/Latinx students. We may wonder how this breakdown would look by major or college as well, which we'll explore later in the lab.
</div>

### Question 5.1

Using the same method but changing the `descending` parameter to `False`, we can discover the lowest fall headcounts. Fill in the following code using the `sort` method to accomplish this.

*Note:* Be sure to use the `ucr_relabeled` table for this question!

```python
ucr_sorted_lowest = ...
ucr_sorted_lowest.show(5)
```

```python
grader.check("q5_1")
```

<div class="alert alert-warning">
The results from this call may not be as interesting as the previous one: we notice that there's a greater breakdown of full-time statuses and new/continuing statuses among these students, and each of these categories only had one individual in the 2022 fall semester. However, if we were to take a look at the next twenty or thirty lowest rows, this may yield some more valuable information about which groups of students are witnessing much lower admission rates.
</div>

---
### Row selection: `where` and the `are` Predicates

The last table method we will talk about is the `where` method. The `where` method keeps all rows that satisfiy a particular boolean condition. It takes in a column label and an `are` statement, which can be crafted using the `are` library. These are the most important `are` library predicates, but there are many more if you would like to investigate: [explore the 'are' library here.](https://data6.org/su25/reference/#tablewhere-predicates)

| Method | Input Type | Method Description |
| --- | --- | --- |
| `are.equal_to(n)` | number | Is the value from the column equal to `n`? |
| `are.above(n)` | number | Is the value from the column above `n`? |
| `are.above_or_equal_to(n)` | number | Is the value from the column above or equal to `n`? |
| `are.below(n)` | number | Is the value from the column below `n`? |
| `are.below_or_equal_to(n)` | number | Is the value from the column below or equal `n`? |
| `are.containing(s)` | string | Is `s` contained in the string value from the given column? |
| `are.containined_in(s)` | string | Is the string value from the given column contained in `s`? |

Adding a `not_` in front of all of these predicates makes each method do the opposite of what it does (ex: `are.not_equal_to(n)`).

*Note*: As we've seen in lecture, we can achieve an **exact match** by not explicitly using an `are` predicate. That is, `where("col", are.equal_to("something")` is identical to `where("col", "something")`; the latter is shorthand for the former.

For example, if we only wanted to see the rows in which the `"Fall Headcount"` was above 500, we could use the `where()` method combined with an `are` method:

```python
ucr_relabeled.where("Fall Headcount", are.above(500))
```

<div class="alert alert-warning">
In the context of our exploration of student demographics, this method can be incredibly useful. For example, if we wanted to delve further into fall headcounts that were above 3 (to avoid getting tons of random rows like we did when using the <code>sort</code> method) but below 20, in order to find groups of students that had relatively low values for fall headcounts, we could use the <code>are.strictly_between</code> predicate:
</div>

```python
ucr_relabeled.where("Fall Headcount", are.strictly_between(1, 21))
```

<div class="alert alert-warning">
From this view of our data, we might be able to see more interesting trends, such as the low values of full-time Native Hawaiian / Pacific Islander undergrads in 2022. We could also chain together multiple <code>where</code> calls in order to satisfy conditions for different rows in our table. If we wanted to get even more specific in looking at these lower headcounts, we could specify the values as we did previously, but then also look for only full-time students:
</div>

```python
ucr_relabeled.where("Fall Headcount", are.strictly_between(1, 21)).where("Full-Time Status", are.equal_to("Full-Time"))
```

<div class="alert alert-warning">
You may notice that some of the rows appear to be duplicates, but they are different -- we just can't see it because we dropped the <code>"New/Continuing Status"</code> column earlier. Now, when we really limit the data like this, we can see some trends about which races/ethnicities have less students belonging to them. This table still includes both transfers and frosh, undergraduates and graduates, but we can see some consistencies with Native Hawaiian/Pacific Islander, Domestic Unknown, and Black/African American students tending to have lower numbers. This type of more complex use of the <code>where</code> method showcases just how useful it can be.
</div>

In the example above, you can see the use of the predicate `are.equal_to` in order to match strings (you can also match exact numbers this way). You can also use a shorthand of omitting the `are` predicate completely to achieve the same effect. For example, if we wanted to see the rows of only undergrads, we could use the shorthand in this way:

```python
ucr_relabeled.where("Level", "Undergraduate")
```

When we omit the `are` predicate, the function is just using `are.equal_to` behind the scenes as the default option!

### Question 5.2
At this point in our analysis of the data, we have many tools under our belt to take a deeper look into the demographics of gender and race/ethnicity. For the following question, looking only at rows where the `"Year"` is equal to 2022, keep only the female-identifying students, and then sort by descending `"Fall Headcount"` value. Because this question is more complex and requires chaining together of multiple calls, we've provided skeleton code for you.

*Hint:* Use the shorthand in the example above with the `where` methods.

```python
ucr_f_2022 = ucr_relabeled.where(..., ...).where(..., ...).sort(..., ...)
ucr_f_2022.show(5)
```

```python
grader.check("q5_2")
```

<div class="alert alert-warning">
From the result above, you should see some similarities to what we found earlier when looking at the highest fall headcounts: female-identifying Chicanx/Latinx students make up some of the highest headcount values, regardless of <code>"New/Continuing Status"</code> or <code>"Original Entry Type"</code>. Additionally from this view, we can see that the next largest proportion of female-identifying students in 2022 had a race/ethnicity of Asian.
</div>

Now that we've learned a bit of surface-level information about the distribution of race/ethnicity, let's dig a little deeper into what the `"Gender"` column looks like. Utilizing the `np.unique` method that we used in the first section above, we can find the distinct values in the `"Gender"` column, to see which values are used to represent gender in this dataset.

```python
# Just run this cell
np.unique(ucr.column("Gender"))
```

As we can see above, the only two options for representing gender in this dataset are "Female" and "Male". It can often be hard to determine how to represent complex pieces of information (especially when this information is about individual people) as categorized versions in data.

<!-- BEGIN QUESTION -->

### Question 5.3 (Discussion)
Because this categorization of gender may not be the best way to represent this information, what do you think some potential issues with looking only at the female-identifying individuals from our data are? Additionally, what are the implications and potential sources of harm when we have to categorize information like this?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<div class="alert alert-warning">
We can tie your takeaways from this question back to one we asked earlier in the lab, while looking at the unique values available for race/ethnicity. There are certainly some issues with the representation of this demographic as well, as many individuals have to be lumped in together (e.g. students with origins from Middle Eastern or Arab countries having to identify themselves as "White", even when that likely does not accurately capture how they would identify themselves).
</div>

---

## Pets of Data 6

That was a lot of heavy lifting, but now you have some table practice under your belt. Good job!

<img src="cat-weights.jpg" width="50%" alt="photo of a cat under a weight"/>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q2_2 (points: 1)
```
>>> unemployment_totals.num_rows == unemployment_rates.num_rows
True
```
```
>>> unemployment_totals.num_columns == 3
True
```


Tests for q2_3 (points: 1)
```
>>> unemployment_races_ethnicities.num_columns == 6
True
```
```
>>> 'Month' in unemployment_races_ethnicities.labels and 'Year' in unemployment_races_ethnicities.labels
True
```
```
>>> '16 to 19 years old' not in unemployment_races_ethnicities.labels
True
```
```
>>> unemployment_races_ethnicities.num_rows == unemployment_rates.num_rows
True
```


Tests for q2_4 (points: 1)
```
>>> 'Total' in unemployment_rate_highest.labels
True
```
```
>>> 'Month' in unemployment_rate_highest.labels
True
```
```
>>> 'Year' in unemployment_rate_highest.labels
True
```
```
>>> unemployment_rate_highest.num_rows == unemployment_rates.num_rows
True
```
```
>>> round(unemployment_rate_highest.row(0).item('Total'), 2) == 14.8
True
```
```
>>> unemployment_rate_highest.row(0).item('Month') == 4
True
```
```
>>> unemployment_rate_highest.row(0).item('Year') == 2020
True
```


Tests for q2_5 (points: 1)
```
>>> all(unemployment_2025.column('Year') == 2025)
True
```
```
>>> unemployment_2025.num_rows == 8
True
```


Tests for q3_1 (points: 1)
```
>>> len(months) == len(years) == unemployment_rates.num_rows
True
```
```
>>> len(months) == len(totals)
True
```


Tests for q3_2 (points: 1)
```
>>> len(times) == len(totals) == len(months)
True
```
```
>>> times.item(5) == 2006
True
```
```
>>> round(times.item(0), 2) == round(2005 + 7 / 12, 2)
True
```
```
>>> round(times.item(-1), 2) == round(2025 + 7 / 12, 2)
True
```


Tests for q3_3 (points: 1)
```
>>> unemployment_over_time.num_columns == 2
True
```
```
>>> unemployment_over_time.row(0).item('Time') == times.item(0)
True
```
```
>>> unemployment_over_time.row(-1).item('Time') == times.item(-1)
True
```
```
>>> unemployment_over_time.row(0).item('Total Unemployment') == totals.item(0)
True
```


Tests for q4_2 (points: [0, 1])
```
>>> assert isinstance(ucr_rows, int)
```
```
>>> assert ucr_rows == 5021
```


Tests for q4_3 (points: [0, 1])
```
>>> assert isinstance(q4_3, np.ndarray)
```
```
>>> assert q4_3 == make_array(3)
```


Tests for q4_4 (points: [0, 0.5, 0.5, 1])
```
>>> assert isinstance(ucr_relabeled, Table)
```
```
>>> assert 'Race/Ethnicity' in ucr_relabeled.labels
```
```
>>> assert 'IPEDS Race/Ethnicity' not in ucr_relabeled.labels
```
```
>>> assert all(ucr_relabeled) == all(ucr.relabeled('IPEDS Race/Ethnicity', 'Race/Ethnicity'))
```


Tests for q4_5 (points: [0, 1, 1])
```
>>> assert isinstance(ucr_selected, Table)
```
```
>>> assert ucr_selected.num_columns == 5
```
```
>>> assert all(ucr_selected) == all(ucr_relabeled.select('Year', 'Level', 'Full-Time Status', 'Race/Ethnicity', 'Fall Headcount'))
```


Tests for q4_6 (points: [0, 1])
```
>>> assert isinstance(q4_6, np.ndarray)
```
```
>>> assert q4_6 == make_array(4)
```


Tests for q5_1 (points: [0, 1, 1, 1])
```
>>> assert isinstance(ucr_sorted_lowest, Table)
```
```
>>> assert ucr_sorted_lowest.take(0).column('Fall Headcount').item(0) == 1
```
```
>>> assert ucr_sorted_lowest.take(1).column('Race/Ethnicity').item(0) == 'Asian'
```
```
>>> assert all(ucr_sorted_lowest) == all(ucr_relabeled.sort('Fall Headcount', descending=False))
```


Tests for q5_2 (points: [0, 1, 1, 1])
```
>>> assert isinstance(ucr_f_2022, Table)
```
```
>>> assert ucr_f_2022.num_rows == 94
```
```
>>> assert all(ucr_f_2022.column('Gender') == 'Female')
```
```
>>> assert all(ucr_f_2022.column('Year') == 2022)
```
