```python
# Initialize Otter
import otter
grader = otter.Notebook("hw03.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Homework 3 – Data Visualization

## Data 6, Fall 2026

This homework is due on **Wednesday, October 14 at 11:59PM**. 

You must submit the assignment to [Gradescope](https://www.gradescope.com/). See the [syllabus](https://data6.org/fa26/syllabus) for our homework late submission policy.

**Reminder:** This homework assignment contains hidden tests, which means that even if the test passed, **it does not mean your answer is 100% correct**. The autograder in the homework is intended as a sanity check.

While we encourage you to collaborate with your peers, directly copying solutions is not allowed. Please refer to our [syllabus](https://data6.org/fa26/syllabus) on Academic Honesty.
If you find yourself stuck on a problem, we recommend that you make a post on [Ed](https://edstem.org) or attend [office hours](https://data6.org/fa26/calendar/).

```python
# just run this cell
from datascience import *
import numpy as np
from IPython.display import YouTubeVideo, HTML, display
from ipywidgets import interact, widgets
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Crimes During July & August 2026 in London

For this assignment, we are going to examine street and search data for crimes in London. This data is publicly available through the UK's open data initiative to be more transparent about searches, locations where crimes have occurred, and the types of crimes. You can find [more information about the data here](https://data.police.uk/).

## Part 1: Locations of Crimes in the City of London

### Question 1a

Read in the `london_street_crime_data.csv` file and assign the data to `street`. We will use `street` for additional analyses. 

```python
street = ...
street
```

```python
grader.check("q1a")
```

### Question 1b

From looking at the data in `street`, what is the unit of analysis, or what does each row of the data represent? Assign `street_unit_of_analysis` to an integer 1, 2, 3, or 4 corresponding to the choices below. This question has a hidden test.

1. Location of crime
2. Crime Victim
3. Officer
4. Suspect


```python
street_unit_of_analysis = ...
```

```python
grader.check("q1b")
```

### Question 1c

Looking at our data, we have several different types of variables. Assign values as strings to each of the variable names below. Your options are `nominal`, `ordinal`, `discrete`, and `continuous`.

This question has hidden tests.


```python
longitude_var = ...
month_var = ...
location_var = ...
crime_type_var = ...

```

```python
grader.check("q1c")
```

<!-- BEGIN QUESTION -->

### Question 1d

Now that we know what type of variables are data are, let's conduct some exploratory data analysis (EDA) to look for patterns in July and August London crimes. Visualizing data can be the best first step to getting a better understanding of our data. First, we'll create a visual to look at the relationship between `longitude` and `latitude`.

For simplicity, we are not overlaying a map of London for visualizations that use `longitude` and `latitude`, but it is possible to incorporate into plots using `matplotlib`.

*Hint*: Longitude corresponds to values that measure east to west positions. Latitude measures values north to south. London is an interesting location because it has the [prime meridian](https://en.wikipedia.org/wiki/Prime_meridian) or the longitude that is defined as $0^\circ$.

<!-- <img src='Outline of London.png' alt= 'Outline of the City of London' width='500' height='300'> -->


```python
...
```

<!-- END QUESTION -->

### Question 1e

This plot doesn't tell us much about where crimes are occurring other than that some points are seemingly outside of the city of London. This could be due to the crime moving outside of the city of London and the city of London's police force still having jurisdiction. 

We'll ignore those points, but if this were a real data analysis you would definitely want to investigate more about the meaining of those crime locations.

For our purposes, let's filter the `street` table to only include values in the general area. Specifically, our longitudes should be between `-0.5` and `0.25` (exclusive). We'll call this new table `london_st`.

```python
london_st = ...
london_st
```

```python
grader.check("q1e")
```

### Question 1f

Now we can look at a scatter plot of the `longitude` and `latitude` using `london_st` to see if our filtered data presents anything of interest about locations of London crimes.

Even when filtering our data and zooming into more specific longitude and latitude values, we can see that most crimes happen in a very dense area of London. This could be due to the dense population of London.

As emerging data scientists, we might then be interested in the types and number of crimes that happen in `london_st`.

Get the counts for each `crime_type` in `london_st`. We'll call this new table `crime_count`.

```python
london_st.scatter('longitude', 'latitude')
```

```python
crime_count = ...
crime_count
```

```python
grader.check("q1f")
```

### Question 1g

We can see a lot of crimes were reported in London during July and August. However, we currently can't see which crimes were the most reported during this period. Let's reorganize our table `crime_count` to see what the most commonly reported crime was in London in July and August. We will call our new table `crime_count_organized`, where the first row shows the crime with the largest count.

The number of each crime is helpful, but it might not be the most impactful to share with people in London who are fearful of how dangerous London is. One way to really show the differences between the types of crimes that are reported is to change the counts to percentages. We are going to add a column `percentage` to our table `crime_count`. In order to get the percentages, we need to get the total number of crimes that have occurred in London over the two months. We will call this value `total_crimes`.

*Hint*: When adding the column `percentage` to `crime_count_organized`, **make sure to change the proportion to a percentage**. We want to make it easier for when we plot the percentages later on.

```python
crime_count_organized = ...
print(crime_count_organized)

total_crimes = ...
print('\n', f'Total number of crimes in July and August is {total_crimes}', '\n')

crime_count_organized_w_pct = ...
crime_count_organized_w_pct

```

```python
grader.check("q1g")
```

<!-- BEGIN QUESTION -->

### Question 1h

Let's visualize the `count` for each `crime_type`. We will use `crime_count_organized_w_pct` because we'll compare the visualization for counts to the visualization that uses percentages. Choose a visualization that's best for showing data with one categorical variable and one numerical variable. Remember that we also want a visualization where we can read each type of crime.

```python
...
plots.title('Number of Crimes in London\nDuring July & August 2026')
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1i

Now let's visualize the `percentage` variable with `crime_type` to see how it can be interpreted differently compared to the visualization that used `count`. We will still use the `crime_count_organized_w_pct` table for this visualization.

```python
...
```

<!-- END QUESTION -->

### Question 1j

Depending on which visualization is shared, they tell our audience slightly different stories. Filter the `crime_count_organized_w_pct` to examine the row with `Robbery` as the crime type. We will save this **table** as `robbery_crime_count_pct`.

```python
robbery_crime_count_pct = ...
robbery_crime_count_pct
```

```python
grader.check("q1j")
```

<!-- BEGIN QUESTION -->

### Question 1k

Let's compare the `count` and `percentage` columns for the `robbery_crime_count_pct` table. If you were looking for an apartment in this area of London, which value would make you feel safer? Why or why not?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 1l

There is also a lot of theft that happens in London. It is important to note that theft is when someone unlawfully takes possessions with the intent to keep, while robbery is when someone takes possessions through threats or force.

Let's create a table with only theft and robbery. You might notice that there is a lot of crime types in `crime_type` that are related to theft, like `Other theft`, `Theft from the person`, `Shoplifting`, and `Bicycle theft`. In the following questions, we'll create a new `theft` variable, but for now we need to filter and restructure our data.

Follow the steps to create a new table, `theft_count`:
1. Using `london_st`, filter the data to only include the theft categories within `crime_type` (`Other theft`, `Theft from the person`, `Shoplifting`, and `Bicycle theft`).
2. Find the counts for each combination of `crime_type` and `last_outcome_category` (From here, you should be able to answer the question, "How many bike thefts are still under investigation?").
3. Reorganize the data so the first row contains the largest count for each crime type and outcome.

```python
theft_count = ...
print(theft_count)

# Just run this code. We'll cover pivot later.
theft_count_wide = theft_count.pivot('last_outcome_category', 'crime_type', 'count', np.sum)
theft_count_wide


```

```python
grader.check("q1l")
```

### Question 1m

Now we have the counts for the specific types of theft, let's remove the columns `Local resolution` and `Under investigation` from the `theft_count_wide` table. We can temporarily save our new table as `theft_count_subset`. 

Then we are going to get the total amounts of `Awaiting court outcome`, `Investigation complete; no suspect identified`, and `Unable to prosecute suspect` to create our new table `only_theft`. 

```python
theft_count_subset = ...
only_theft = Table().with_columns('crime_type', make_array('theft'),
  'Awaiting court outcome', ...,
  'Investigation complete; no suspect identified', ...,
  'Unable to prosecute suspect', ...
  )

only_theft
```

```python
grader.check("q1m")
```

### Question 1n

Let's now filter for only `Burglary` from `crime_type` in the `london_st` table. Similar to the `theft_count` table above, we will want to get the counts of `crime_type` and `last_outcome_category` so our data matches.

Create a one-row table called `burglary_count` that contains the same columns as `only-theft`.

We (the Data 6 staff) will clean up the table to make `theft_burglary_wide` with counts for our new `theft` and `Burglary` types of crime.

```python
burglary_count = ...
burglary_count
```

```python
# only run the code below this comment once
burglary_count_wide = burglary_count.pivot('last_outcome_category', 'crime_type', 'count', np.sum)

burglary_count_subset = burglary_count_wide.drop('Under investigation')

theft_burglary_wide = only_theft.with_row(burglary_count_subset.row(0))
theft_burglary_wide
```

```python
grader.check("q1n")
```

## Part 2 - Police Searches

Now that we have looked at where crimes occurred in the city of London in July and August, let's look at who is being searched by London police. This comes from the same data as the crime location data above.

### Question 2a

Read in the `london_search_crime_data.csv` file and assign the data to `search`. We will use `search` for additional analyses. 

```python
search = ...
search
```

```python
grader.check("q2a")
```

### Question 2b

From that data, we might be interested in studying the differences for each sex (named `gender` in the `search` table). We are only interested in the type of search (`type`) where people are searched on their person (`Person search`). After that, create a table `date_gender_counts` that shows the count of how many person searches occurred for each `gender` on each `date`.

To pass the autograder, `date` should be the first column.

```python
date_gender_counts = ...
date_gender_counts
```

```python
grader.check("q2b")
```

### Question 2c

From our `date_gender_counts` table, we see that males show up more often in our table than females. Let's see how many person searches there were **on average** for males and females over the entire two months. 
Remember that to get calculations for multiple groups for our `gender` variable, we need to add the function that we want to calculate for our `gender` variable. 

```python
avg_search_by_gender = ...

avg_search_by_gender.barh('gender', 'count mean')
plots.title('Average Number of Searches\non Males & Females')
```

```python
grader.check("q2c")
```

<!-- BEGIN QUESTION -->

### Question 2d

Having the averages is useful for identifying the pattern that males are searched more often than females on average. However, our bar chart does not tell us the number of searches over time. Since we know that males were searched more over the period of our data collection, we are going to look at the counts of males being searched over time using the `date_gender_counts` table by following the steps below.

1. Let's filter the data to only include `Male` for `gender`.
2. Then choose a visualization that's best for showing data over time.

Your solution should utilize method chaining and be one line of code.

*Hint*: Your x-axis will look indiscernible. We'll learn to fix this later in the semester.

```python
...
```

<!-- END QUESTION -->

### Question 2e

Looking at the `search` table, we once again have several different types of variables. Assign values as strings to each of the variable names below. Your options are `nominal`, `ordinal`, `discrete`, and `continuous`.

This question has hidden tests.

```python
search_age_range = ...
search_self_defined_ethnicity = ...
search_object_of_search = ...
search_outcome = ...
```

```python
grader.check("q2e")
```

## Pets of Data 6

Sunkist needs some head scratching before going to sleep. Can you help after submitting to Gradescope?

<img src="sunkist.jpg" width="50%" alt="cat lying down"/>

## Submission

Below, you will see two cells. Running the first cell will automatically generate a PDF of all questions that need to be manually graded, and running the second cell will automatically generate a zip with your autograded answers. You are responsible for submitting both the coding portion (the zip) and the written portion (the PDF) to their respective Gradescope portals. **Please save before exporting!**

> **Important: You must correctly assign the pages of your PDF after you submit to the correct gradescope assignment. If your pages are not correctly assigned and/or not in the correct PDF format by the deadline, we reserve the right to award no points for your written work.**

If there are issues with automatically generating the PDF in the first cell, you can try downloading the notebook as a PDF by clicking on `File -> Save and Export Notebook As... -> PDF`. If that doesn't work either, you can manually take screenshots of your answers to the manually graded questions and submit those. Either way, **you are responsible for ensuring your submission follows our requirements, we will NOT be granting regrade requests for submissions that don't follow instructions.**

> **TO AVOID AUTOGRADER ERRORS: Before running the cells below, _SAVE_ your work (either with the keyboard shortcut `CMD+S` for Macs, or `Ctrl+S` for Windows, or by clicking on the floppy disk in the top left) and click `Run -> Run All Cells` in the toolbar.**

```python
from otter.export import export_notebook
from os import path
from IPython.display import display, HTML
name = 'project1'
export_notebook(f"{name}.ipynb", filtering=True, pagebreaks=True, exporter_type = "html")
if(path.exists(f'{name}.pdf')):
    display(HTML(f"Download your PDF <a href='{name}.pdf' download>here</a>."))
else:
    print("\n Pdf generation failed, please try the other methods described.")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1a (points: 1)
```
>>> type(street) == Table
True
```
```
>>> street.num_rows
1398
```


Tests for q1b (points: 1)
```
>>> street_unit_of_analysis in [1, 2, 3, 4]
True
```


Tests for q1c (points: [0])
```
>>> [x in ['nominal', 'ordinal', 'discrete', 'continuous'] for x in [longitude_var, month_var, location_var, crime_type_var]]
[True, True, True, True]
```


Tests for q1e (points: 1)
```
>>> expect_london_st = street.where('longitude', are.strictly_between(-0.5, 0.25))
>>> london_st.num_rows == expect_london_st.num_rows
True
```


Tests for q1f (points: 1)
```
>>> expect_crime_count = london_st.group('crime_type')
>>> crime_count.num_rows == expect_crime_count.num_rows
True
```


Tests for q1g (points: [1, 1, 1])
```
>>> expect_crime_count_organized = crime_count.sort('count', descending=True)
>>> crime_count_organized.column('crime_type').item(0) == expect_crime_count_organized.column('crime_type').item(0)
True
```
```
>>> total_crimes == 1393
True
```
```
>>> np.isclose(crime_count_organized_w_pct.where('crime_type', 'Other theft').column('percentage').item(0), 22.7566)
True
```


Tests for q1j (points: [0, 0])
```
>>> robbery_crime_count_pct.labels
('crime_type', 'count', 'percentage')
```
```
>>> robbery_crime_count_pct.num_rows == 1
True
```


Tests for q1l (points: 1)
```
>>> theft_count.column('count').item(3) == 85
True
```
```
>>> theft_count.where('last_outcome_category', 'Under investigation').num_rows == 4
True
```


Tests for q1m (points: [0.5, 0.5, 1, 1, 1])
```
>>> expect_theft_count_subset = theft_count_wide.drop('Local resolution', 'Under investigation')
>>> theft_count_subset.num_columns == expect_theft_count_subset.num_columns
True
```
```
>>> only_theft.num_rows == 1
True
```
```
>>> sum(theft_count_subset.column('Awaiting court outcome')) == 13
True
```
```
>>> sum(theft_count_subset.column('Investigation complete; no suspect identified')) == 413
True
```
```
>>> sum(theft_count_subset.column('Unable to prosecute suspect')) == 62
True
```


Tests for q1n (points: 1)
```
>>> burglary_count.where('last_outcome_category', 'Under investigation').column('count').item(0)
59
```


Tests for q2a (points: 1)
```
>>> expect_search = Table.read_table('london_search_crime_data.csv')
>>> expect_search.num_rows == search.num_rows
True
```


Tests for q2b (points: 2)
```
>>> date_gender_counts.labels
('date', 'gender', 'count')
```
```
>>> date_gender_counts.where('date', '2026-07-12').column('count').item(0) == 2
True
```


Tests for q2c (points: [0])
```
>>> np.isclose(avg_search_by_gender.where('gender', 'Male').column('count mean').item(0), 3.3617021276595747)
True
```


Tests for q2e (points: [0])
```
>>> [x in ['nominal', 'ordinal', 'discrete', 'continuous'] for x in [search_age_range, search_self_defined_ethnicity, search_object_of_search, search_outcome]]
[True, True, True, True]
```
