```python
# Initialize Otter
import otter
grader = otter.Notebook("lab07.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Lab 7 – Food Safety

## Data 6, Fall 2026
For this week, we're going to spend more time exploring the food safety data from the San Francisco Department of Public Health. We will continue to interpret data and utilize visualizations to draw conclusions from. 

This lab is due on **Sunday, October 11th at 11:59PM**.

You must submit the assignment to [Gradescope](https://www.gradescope.com/).

```python
# Run this cell to load all required Python libraries
import numpy as np
from datascience import *

# import seaborn as sns
import matplotlib.pyplot as plt
plt.style.use("fivethirtyeight")
%matplotlib inline
```

## San Francisco (SF) Department of Public Health's Food Safety Data

While examining the data through Google Sheets was useful to see specific rows and adding filters and the visualizations were helpful for creating one-time plots, having the capability of doing both together makes using a notebook a better alternative for data manipulation and visualization.

**_Note_**: There might be slight discrepancies between this dataset and the data you used from the SF Department of Public Health because this data was downloaded and cleaned a few days ago.

### Question 1a

Read in the `Health_Inspection_scores.csv` file and assign the data to `safe`. We will use `safe` for additional analyses. 

```python
safe = ...
safe
```

```python
grader.check("1a")
```

<!-- BEGIN QUESTION -->

### Question 1b

From looking at the data during the discussion, you might have seen the varying times for the total amount of time inspectors spend conducting each inspection (`total_time`). Let's use the `safe` table to visualize the association between `violation_count` and `total_time`. 

*Hint*: We have already cleaned the data partially by removing the negative values from the data. In a real scenario, we would most likely have to reach out to the maintainer of the data to see if those negative values are typos or if they correspond to some internal way of coding the data.

```python
...
plt.xlabel('Number of violations')
plt.ylabel('Total Time of Inspection')
```

<!-- END QUESTION -->

### Question 1c

From looking at the data, it is interesting to see the group of inspections that are separated from the rest of the data. Let's investigate this data further by creating a new table, `safe_long_total_time`, where we filter the data to only focus on the long total times for inspections. For our purposes, let's filter `total_time` for values that are **greater than 400**. 

```python
safe_long_total_time = ...
safe_long_total_time
```

```python
grader.check("1c")
```

### Question 1d

With this subset of data (`safe_long_total_time`), let's investigate what types of establishments are being inspected for such long periods of time. From our data, choose the variable where we can count the number of those types of establishments being inspected. Then we should reorder our data to have the greatest number of **establishment types** at the top of our new table `safe_long_total_time_organize`.

```python
safe_long_total_time_organize = ...
safe_long_total_time_organize
```

```python
grader.check("1d")
```

### Question 1e

It seems like restaurants are the most likely to be inspected for longer periods of time. Let's explore a little further into these long inspections. Within the `safe_long_total_time` table, we could examine other potential confounding variables that may be responsible for the long inspection times. 

Let's examine whether there were differences in the number of inspections by `inspector`. Let's also reorganize the data to have the first row be the greatest number of inspections from a specific inspectore. Our new table will be called `inspector_counts`.

Let's then only focus on who the top 5 inspectors are. We'll assign these 5 inspectors in the table `top5_inspectors`.

Finally, let's make an ordinal variable that ranks these inspectors from 1 (greater number of long inspections) to 5 (least number of long inspections). We'll called this table with rankings `top5_ranked_inspectors`.


```python
inspector_counts = ...
print(inspector_counts)

top5_inspectors = ...
print(top5_inspectors)

top5_ranked_inspectors = ...
top5_ranked_inspectors
```

```python
grader.check("1e")
```

<!-- BEGIN QUESTION -->

### Question 1f

From the top 5 inspectors, we can see that Patrick Wood has many more long inspection times in comparison to their coworkers. Let's analyze data on Patrick Wood to investigate why they have so many long inspections.

We will use the `safe` table to filter for inspections completed by Patrick Wood. We'll save the table with Patrick's inspections in `patrick`.

```python
patrick = ...
patrick
```

```python
grader.check("1f")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1g

We've examined long inspections with all the inspectors' data, but let's create a visualization looking at the association between `violation_count` and `total_time` for Patrick Wood.

```python
# Calculates the average while ignoring the nan values
print(np.nanmean(patrick.column('violation_count')))

...
plt.title('Violation Count & Total Time\nAssociation for Patrick Wood')
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1h

Comparing the scatter plot from Question 1b to the scatter plot for Patrick Wood, what comparisons or differences can we make between Patrick Wood's data and the full dataset?

Why does Patrick Wood have so many nan (not a number) values in their violation count?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Question 2

We've completed some exploratory data analysis further examining why some inspectors take longer to complete some inspections and we've examined that the longer inspection times have mainly been for various types of restaurants, but now let's focus more attention on differences between locations that pass, conditionally pass, or fail their inspections.

### Question 2a

Let's now explore if there are differences in the average `total_time` and `violation_count` variables by 'facility_rating_status'. We will assign these calculations to the table `status_avg_time_violation`. 

*Hint*: Make sure your `status_avg_time_violation` table only has three columns.

```python
status_avg_time_violation = ...(
  ...
  )...(
    ..., np.nanmean
    )

status_avg_time_violation
```

```python
grader.check("2a")
```

<!-- BEGIN QUESTION -->

### Question 2b

How is it possible that there are closures and conditional passes that have the same average number of violations, but result in two different outcomes for locations?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 2c

Let's now explore if there are differences in the average `total_time` and `violation_count` variables by `facility_rating_status` and `inspector`. We will use the `safe` table to create a new table called `avg_time_violation_status_inspector`. 

*Hint*: Make sure your `avg_time_violation_status_inspector` only has four columns.

```python
avg_time_violation_status_inspector = ...(
  ...
  )...(
    ..., np.nanmean
    )

avg_time_violation_status_inspector
```

```python
grader.check("2c")
```

### Question 2d

From question 1, we might be interested to know more about where Patrick Wood is in `avg_time_violation_status_inspector`. 

Let's filter `avg_time_violation_status_inspector` to see what Patrick Wood's average total time for inspections and average violation count for each `facility_rating_status`.

We'll call the table with only Patrick Wood's data `patrick_avg_time_violation`.

```python
patrick_avg_time_violation = ...
patrick_avg_time_violation
```

```python
grader.check("2d")
```

<!-- BEGIN QUESTION -->

### Question 2e

What can we interpret from Patrick Wood's `total_time nanmean` and `violation_count nanmean` values for each `facility_rating_status`?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2f

Let's further explore the `avg_time_violation_status_inspector` through visualizations. First, let's create a visualization that examines the `total_time nanmean` for each inspector, but only for Closure `facility_rating_status`.

```python
...
plt.title('Average Total Time for Inspections\nBy Inspector')
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2g

While still using the `avg_time_violation_status_inspector` table and while continuing to filter for Closure `facility_rating_status`, let's examine the average number of violation counts for each inspector using the appropriate visualization.

```python
...
plt.title('Average Number of Violations\nfor Facilities That Needed to Close\nby Inspector')
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2h

Finally, we will still use the `avg_time_violation_status_inspector` table, but let's now filter for `Conditional Pass` and visualize the relationship between `inspector` and `violation_count nanmean` to compare the average violation counts per inspector for `Conditional Pass` and the plot above for `Closure` `facility_rating_status`.

```python
...
plt.title('Average Number of Violations\nfor Conditional Passing Facilities\nby Inspector')
```

<!-- END QUESTION -->

## Pets of Data 6

Lucy waits for no one when it is time for bed. She is ready for you submit your lab to Gradescope and call it a night.

<img src="lucy.png" width="50%" alt="cat laying down as a muffin on a bed"/>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for 1a (points: 1)
```
>>> safe.num_rows == 22333
True
```


Tests for 1c (points: [0, 1])
```
>>> safe_long_total_time.column('inspector').item(1) == 'Patrick Wood'
True
```
```
>>> safe_long_total_time.num_rows == 84
True
```


Tests for 1d (points: 1)
```
>>> safe_long_total_time_organize.column('count').item(1) == 11
True
```


Tests for 1e (points: 3)
```
>>> inspector_counts.labels
('inspector', 'count')
```
```
>>> top5_inspectors.column('count').item(0) == 18
True
```
```
>>> top5_ranked_inspectors.column('most_long_inspections').item(0) == 1
True
```


Tests for 1f (points: 1)
```
>>> np.unique(patrick['inspector']).item(0) == 'Patrick Wood'
True
```


Tests for 2a (points: [0, 1])
```
>>> status_avg_time_violation.num_columns == 3
True
```
```
>>> np.isclose(status_avg_time_violation.column('violation_count nanmean').item(0), 7.1704)
True
```


Tests for 2c (points: 2)
```
>>> avg_time_violation_status_inspector.num_columns == 4
True
```
```
>>> avg_time_violation_status_inspector.where('facility_rating_status', 'nan').where('inspector', 'Amelia Castelli').column('violation_count nanmean').item(0) == 1
True
```


Tests for 2d (points: 1)
```
>>> patrick_avg_time_violation.where('facility_rating_status', 'Closure').column('violation_count nanmean').item(0) == 4
True
```
